## Data structures

Data Structures in this tutorial
1. Tuples
2. Dictionaries
3. Arrays

<br>
As an overview, tuples and arrays are both ordered sequences of elements (so we can index into them). Dictionaries and arrays are both mutable.

### 1.0 Tuples

We can create a tuple by enclosing an ordered collection of elements in `( )`.

Syntax: <br>
```(item1, item2, ...)```

```julia
myfavoriteanimals = ("owls", "cats", "mantis shrimp")
```

Output:
```
("owls", "cats", "mantis shrimp")
```

We can index into this tuple,

```julia
myfavoriteanimals[1]
```

Output:
```
"owls"
```

but since tuples are immutable, we can't update it

```julia
myfavoriteanimals[1] = "otters"
```

Output:
```
Stacktrace:
 [1] top-level scope at In[3]:1
 [2] include_string(::Function, ::Module, ::String, ::String) at .\loading.jl:1091
```

### 1.1 NamedTuples

`NamedTuple`s are just like `Tuple`s except that each element additionally has a name! They have a special syntax using `=` inside a tuple:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals = (bird = "owl", mammal = "cats", sea_creature = "mantis shrimp")
```

Output:
```
(bird = "owl", mammal = "cats", sea_creature = "mantis shrimp")
```

Like regular `Tuples`, `NamedTuples` are ordered, so we can retrieve their elements via indexing:

```julia
myfavoriteanimals[1]
```

Output:
```
"owl"
```

They also add the special ability to access values by their name:

```julia
myfavoriteanimals.bird
```

Output:
```
"owl"
```

### 2. Dictionaries

If we have sets of data related to one another, we may choose to store that data in a dictionary. We can create a dictionary using the `Dict()` function, which we can initialize as an empty dictionary or one storing key, value pairs.

Syntax:
``` Dict(key1 => value1, key2 => value2, ...) ```


```julia
myphonebook = Dict("Jenny" => "867-5309", "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String,String} with 2 entries:
  "Jenny"        => "867-5309"
  "Ghostbusters" => "555-2368"
```

In this example, each name and number is a "key" and "value" pair. We can grab Jenny's number (a value) using the associated key

```julia
myphonebook["Jenny"]
```

Output:
```
"867-5309"
```

We can add another entry to this dictionary as follows

```julia
myphonebook["Kramer"] = "555-FILK"
```

Output:
```
"555-FILK"
```

Let's check what our phonebook looks like now...<br>
Note: Pay attention to the order 

```julia
myphonebook
```

Output:
```
Dict{String,String} with 3 entries:
  "Jenny"        => "867-5309"
  "Kramer"       => "555-FILK"
  "Ghostbusters" => "555-2368"
```

We can **delete** Kramer from our contact list - and simultaneously grab his number - by using `pop!`

```julia
pop!(myphonebook, "Kramer")
```

Output:
```
"555-FILK"
```

```julia
myphonebook
```

Output:
```
Dict{String,String} with 2 entries:
  "Jenny"        => "867-5309"
  "Ghostbusters" => "555-2368"
```

Unlike tuples and arrays, dictionaries are not ordered. So, we **cannot index into them.**

```julia
myphonebook[1]
```

Output:
```
[1] getindex(::Dict{String,String}, ::Int64) at .\dict.jl:467
 [2] top-level scope at In[14]:1
 [3] include_string(::Function, ::Module, ::String, ::String) at .\loading.jl:1091
```

In the example above, `julia` thinks you're trying to access a value associated with the key `1`.

### 3. Arrays

Unlike tuples, arrays are mutable. Unlike dictionaries, arrays contain ordered collections. <br>
We can create an array by enclosing this collection in `[ ]`.

Syntax: <br>
``` [item1, item2, ...]```


For example, we might create an array to keep track of my friends

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]
```

Output:
```
5-element Array{String,1}:
 "Ted"
 "Robyn"
 "Barney"
 "Lily"
 "Marshall"
```

The `1` in `Array{String,1}` means this is a one dimensional vector.  An `Array{String,2}` would be a 2d matrix, etc.  The `String` is the type of each element. Format: {Type, dimension}

or to store a sequence of numbers

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
mixture = [1, 1, 2, 3, "Ted", "Robyn"]
```

Output:
```
6-element Array{Any,1}:
 1
 1
 2
 3
  "Ted"
  "Robyn"
```

Once we have an array, we can grab individual pieces of data from inside that array by indexing into the array. For example, if we want the third friend listed in `myfriends`, we write

```julia
myfriends[3]
```

Output:
```
"Barney"
```

We can use indexing to edit an existing element of an array

```julia
myfriends[3] = "Broseph"
```

Output:
```
"Broseph"
```

Yes, Julia is **1-based indexing**, not 0-based like Python.  Wars are fought over lesser issues. 

We can also edit the array by using the `push!` and `pop!` functions. `push!` adds an element to the end of an array and `pop!` removes the last element of an array.

We can add another number to our fibonnaci sequence

```julia
push!(fibonacci, 21)
```

Output:
```
8-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 13
 21
```

and then remove it

```julia
pop!(fibonacci)
```

Output:
```
21
```

```julia
fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 13
```

So far I've given examples of only 1D arrays of scalars, but arrays can have an arbitrary number of dimensions and can also store other arrays. 
<br><br>
For example, the following are arrays of arrays:

```julia
favorites = [["koobideh", "chocolate", "eggs"],["penguins", "cats", "sugargliders"]]
```

Output:
```
2-element Array{Array{String,1},1}:
 ["koobideh", "chocolate", "eggs"]
 ["penguins", "cats", "sugargliders"]
```

```julia
## Jagged Array
numbers = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

Output:
```
3-element Array{Array{Int64,1},1}:
 [1, 2, 3]
 [4, 5]
 [6, 7, 8, 9]
```

Below are examples of 2D and 3D arrays populated with **random values** using `rand(dim1, dim2,....)`.

```julia
rand(4, 3)
```

Output:
```
4×3 Array{Float64,2}:
 0.190651  0.25347   0.0141378
 0.440755  0.476866  0.363689
 0.285175  0.516562  0.666533
 0.765735  0.159144  0.405553
```

```julia
rand(4, 3, 2)
```

Output:
```
4×3×2 Array{Float64,3}:
[:, :, 1] =
 0.41674   0.577817   0.29565
 0.315887  0.517629   0.16488
 0.778065  0.0354329  0.602505
 0.188135  0.389743   0.0589714

[:, :, 2] =
 0.440066  0.586505  0.608194
 0.333615  0.612166  0.635925
 0.325759  0.42221   0.55601
 0.266801  0.313602  0.467883
```

Be careful when you want to copy arrays!

```julia
fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somenumbers = fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somenumbers[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Array{Int64,1}:
 404
   1
   2
   3
   5
   8
  13
```

Editing `somenumbers` caused `fibonacci` to get updated as well!

In the above example, we didn't actually make a copy of `fibonacci`. We just created a new way to access the entries in the array bound to `fibonacci`. **Makes a shallow copy if = used.**

If we'd like to make a copy of the array bound to `fibonacci`, we can use the `copy` function.

```julia
# First, restore fibonacci
fibonacci[1] = 1
fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 13
```

Use `copy()` to make a **deep copy**

```julia
somemorenumbers = copy(fibonacci)
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somemorenumbers[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 13
```

In this last example, fibonacci was not updated. Therefore we see that the arrays bound to `somemorenumbers` and `fibonacci` are distinct.

### Exercises from the tutorial series 

#### 3.1 
Create an array, `a_ray`, with the following code:

```a_ray = [1, 2, 3]```

Add the number `4` to the end of this array and then remove it.

```julia
a_ray = [1,2,3]
```

Output:
```
3-element Array{Int64,1}:
 1
 2
 3
```

```julia
push!(a_ray,4)
```

Output:
```
4-element Array{Int64,1}:
 1
 2
 3
 4
```

```julia
pop!(a_ray)
```

Output:
```
4
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Try to add "Emergency" as key to `myphonebook` with the value `string(911)` with the following code
```julia
myphonebook["Emergency"] = 911
```

Why doesn't this work?

```julia
myphonebook["Emergency"] = string(911)
```

Output:
```
"911"
```

#### 3.3 
Create a new dictionary called `flexible_phonebook` that has Jenny's number stored as an integer and Ghostbusters' number stored as a string with the following code

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

```julia
## Key value pair, value does not have a pre defined type 
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String,Any} with 2 entries:
  "Jenny"        => 8675309
  "Ghostbusters" => "555-2368"
```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 3.4 
Add the key "Emergency" with the value `911` (an integer) to `flexible_phonebook`.

```julia
flexible_phonebook["Emergency"] = 911
```

Output:
```
911
```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

#### 3.5 
Why can we add an integer as a value to `flexible_phonebook` but not `myphonebook`? How could we have initialized `myphonebook` so that it would accept integers as values?

```julia
myphonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String,Any} with 2 entries:
  "Jenny"        => 8675309
  "Ghostbusters" => "555-2368"
```

Please click on `Validate` button at the top, upon completion of the exercise