# Basic linear algebra in Julia
Author: Andreas Noack Jensen (MIT) (http://www.econ.ku.dk/phdstudent/noack/)
(with edits from Jane Herriman)

Julia provides native implementations of many common and useful linear algebra operations which can be loaded with using LinearAlgebra. Basic operations, such as [tr](https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/#LinearAlgebra.tr) (trace), det (determinant), and inv (inverse) are all supported.
[Linear Algebra in Julia](https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/)

First let's define a random matrix

```julia
A = rand(1:4,3,3)
```

Output:
```
3×3 Array{Int64,2}:
 3  1  3
 1  1  2
 1  2  2
```

```julia
using LinearAlgebra ## ***
println("Trace is: ", tr(A))
println("Determinant is: ",det(A))
println("Inverse Matrix: ",inv(A))
```

Output:
```
Trace is: 6
Determinant is: -3.0000000000000004
Inverse Matrix: [0.6666666666666666 -1.3333333333333333 0.33333333333333337; -5.551115123125783e-17 -0.9999999999999999 1.0; -0.33333333333333326 1.6666666666666665 -0.6666666666666666]
```

```julia
eigvals(A)
eigvecs(A)
```

Output:
```
3×3 Array{Float64,2}:
 -0.392251  -0.860543  -0.783373
 -0.63338    0.240644  -0.400353
  0.66706    0.44895   -0.475441
```

Define a vector of ones

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Output:
```
3-element Array{Float64,1}:
 1.0
 1.0
 1.0
```

Notice that $A$ has type Array{Int64,2} but $x$ has type Array{Float64,1}. Julia defines the aliases Vector{Type}=Array{Type,1} and Matrix{Type}=Array{Type,2}. 

Many of the basic operations are the same as in other languages
#### Multiplication

```julia
b = A*x
```

Output:
```
3-element Array{Float64,1}:
 8.0
 7.0
 8.0
```

#### Transposition
As in other languages `A'` is the conjugate transpose, or adjoint

```julia
A'
```

Output:
```
3×3 LinearAlgebra.Adjoint{Int64,Array{Int64,2}}:
 2  3  4
 3  3  1
 3  1  3
```

and we can get the transpose with

```julia
transpose(A)
```

Output:
```
3×3 LinearAlgebra.Transpose{Int64,Array{Int64,2}}:
 2  3  4
 3  3  1
 3  1  3
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

Output:
```
3×3 Array{Int64,2}:
 29  19  21
 19  19  15
 21  15  19
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ function.

```julia
A\b
```

Output:
```
3-element Array{Float64,1}:
 1.0
 1.0
 1.0
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Array{Float64,2}:
 0.88576   0.373489
 0.811966  0.670621
 0.590453  0.123479
```

```julia
Atall\b
```

Output:
```
2-element Array{Float64,1}:
 12.764367511919707
 -5.680539768328174
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
rankdef = hcat(v, v)
```

Output:
```
3×2 Array{Float64,2}:
 0.608957  0.608957
 0.777304  0.777304
 0.227995  0.227995
```

```julia
rankdef\b
```

Output:
```
2-element Array{Float64,1}:
 5.90876633795577
 5.908766337955771
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Array{Float64,2}:
 0.222785   0.657824  0.903153
 0.0215091  0.668615  0.723462
```

```julia
Ashort\bshort
```

Output:
```
3-element Array{Float64,1}:
 0.6828012725411605
 0.0845169997791721
 0.7369389323198692
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`.



```julia
v = [1,2,3]
```

Output:
```
3-element Array{Int64,1}:
 1
 2
 3
```

```julia
dot_v = sum(v.*v)
# dot_v = v'v
```

Output:
```
14
```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `outer_v`

```julia
outer_v = v*v'
cross_v = [0,0,0]
```

Output:
```
3-element Array{Int64,1}:
 0
 0
 0
```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```

Please click on `Validate` on the top, once you are done with the exercises.